# Lab 01: CSS 120

In this self-paced lab, you'll get hands-on practice with the following **concepts**:

- Open file
- Perform Exploratory Data Analysis

### Note on hidden tests

Some of the questions in this lab will have **hidden tests**. This means that in addition to the `assert` statements you see in the code, there are *hidden* tests. The point is to check whether your code **generalizes** to other cases.

The lab will notify you when/where there is a hidden test, even if you can't see it.

### Motivation

One issue with climate change mitigation policies is that people prefer them but not close to their homes. For instance, imagine that SDGE intends to put a large solar farm close to your place. Imagine furthermore that the mountain you love to hike will have a solar farm. It seems something not very cool indeed.

In this Lab we are going to investigate a dataset of wind turbine proposals and government party vote share in Canada.

## Question 01

Load pandas and numpy using their common aliases.

<!-- grade_id: cell-d39a0941248e0e9c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import pandas as pd
import numpy as np
### END SOLUTION
```

<!-- grade_id: cell-33daa3a69d2b124b, points: 1, grade: true, solution: false -->
```python
assert pd and np
```

## Question 02

The dataset is in <https://raw.githubusercontent.com/umbertomig/CSS120/main/stokes_electoral_2015.csv>. Load it on pandas, saving it as `dat`.

<!-- grade_id: cell-a6ac96a88373be65, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
dat = pd.read_csv("https://raw.githubusercontent.com/umbertomig/CSS120/main/stokes_electoral_2015.csv")
### END SOLUTION
```

<!-- grade_id: cell-46c1fb7f4e9f0a8d, points: 1, grade: true, solution: false -->
```python
assert dat.shape[0] == 1500 and type(dat) == pd.DataFrame
```

## Question 03

What is the average vote share of the Liberal Party (the ruling party at the time)?

1. Save the results as `avg_libvs`.
2. The `perc_lib` variable holds this information.

```python
avg_libvs = ...
```

<!-- grade_id: cell-a6926ce66023eff4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
avg_libvs = dat.perc_lib.mean()
### END SOLUTION
```

<!-- grade_id: cell-337a7390947f3aa5, points: 1, grade: true, solution: false -->
```python
assert float(str(avg_libvs)[:4]) < 1/3 and float(str(avg_libvs)[:5]) > 1/3
```

## Question 04

How many turbine projects were proposed?

1. Save the results as `nprojs`.
2. The `propturbine` variable holds this information.

```python
nprojs = ...
```

<!-- grade_id: cell-660c23dedc0a904b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
nprojs = dat.propturbine.sum()
### END SOLUTION
```

<!-- grade_id: cell-21bdfd0159d9c9ff, points: 1, grade: true, solution: false -->
```python
assert nprojs % 42 == nprojs // 42
```

## Question 05

How many projects were proposed in 2003?

1. Save the results as `nprojs2003`.
2. The `propturbine` variable holds this information.

```python
nprojs2003 = ...
```

<!-- grade_id: cell-40c324465401933f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
nprojs2003 = dat.loc[dat.elecyear == 2003].propturbine.sum()
### END SOLUTION
```

<!-- grade_id: cell-394b37fc7f1d8a22, points: 1, grade: true, solution: false -->
```python
assert not nprojs2003 % 2
```

## Question 06

How many projects were proposed in 2007?

1. Save the results as `nprojs2007`.
2. The `propturbine` variable holds this information.

```python
nprojs2007 = ...
```

<!-- grade_id: cell-0d037aa77e24b657, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
nprojs2007 = dat.loc[dat.elecyear == 2007].propturbine.sum()
### END SOLUTION
```

<!-- grade_id: cell-28d2b21695542522, points: 1, grade: true, solution: false -->
```python
assert (nprojs2007 % 7) / (nprojs2007 // 3) <= 4/3
```

## Question 07

How many projects were proposed in 2011?

1. Save the results as `nprojs2011`.
2. The `propturbine` variable holds this information.

```python
nprojs2011 = ...
```

<!-- grade_id: cell-08b37e4fab67f12b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
nprojs2011 = dat.loc[dat.elecyear == 2011].propturbine.sum()
### END SOLUTION
```

<!-- grade_id: cell-1509676154845be9, points: 1, grade: true, solution: false -->
```python
assert not nprojs2011 % 2 and nprojs2011 > nprojs2007
```

## Question 08

What was the vote share in the places with proposals?

1. Save the results as `vs_prop`.
2. The `propturbine` and `perc_lib` variables hold these informations.

```python
vs_prop = ...
```

<!-- grade_id: cell-910302ae9c152cc4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
vs_prop = dat.loc[dat.propturbine].perc_lib.mean()
### END SOLUTION
```

<!-- grade_id: cell-6182a271aa05f41a, points: 1, grade: true, solution: false -->
```python
assert vs_prop / avg_libvs < 1
```

## Question 09

We may think that places with wind turbine had a lower vote share in the future because of the wind turbines. However, it could be that the government placed the wind turbines in places where voters disliked the government. Is that true?

To answer that, look the vote share in 2003 at places that were assigned wind turbines in the future versus places that were not assigned wind turbines.

1. Save the results as `mean_diff_vshare_before`.
2. The `propturbine`, `elecyear`, `precinct`, and `perc_lib` variables hold these informations.

```python
mean_diff_vshare_before = ...
```

<!-- grade_id: cell-4252e4d8a4150382, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mean_diff_vshare_before = dat.loc[(dat.precinct.isin(dat.loc[dat.propturbine == 1].precinct)) & (dat.elecyear == 2003)].perc_lib.mean() - dat.loc[(dat.precinct.isin(dat.loc[dat.propturbine == 0].precinct)) & (dat.elecyear == 2003)].perc_lib.mean()
### END SOLUTION
```

<!-- grade_id: cell-e5dc50d00afbd4d0, points: 1, grade: true, solution: false -->
```python
assert abs(mean_diff_vshare_before) > 1/29 and abs(mean_diff_vshare_before) < 1/28
```

## Question 10

True or False:

1. The vote share in places with turbine proposals is smaller than the overall voteshare (0.25pt)
1. There were more turbine projects in 2003 than in 2007 (0.25pt)
1. There are more turbine projects in 2007 than in 2011 (0.25pt)
1. Looking at 2003, vote shares, places with turbine proposals later had an average vote share higher thanplaces without turbine proposals (0.25pt)

Hint: To answer, make `q10_i = True` or `q10_i = False`

```python
q10_1 = ...
q10_2 = ...
q10_3 = ...
q10_4 = ...
```

<!-- grade_id: cell-c2285ae7f259b54c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q10_1 = True
q10_2 = False
q10_3 = False
q10_4 = True
### END SOLUTION
```

<!-- grade_id: cell-21d297f13fd71d87, points: 0.25, grade: true, solution: false -->
```python
assert q10_1 or q10_2
```

<!-- grade_id: cell-bdbf316f96e28882, points: 0.25, grade: true, solution: false -->
```python
assert q10_3 or q10_4
```

<!-- grade_id: cell-55ebc2ba5cbecf8e, points: 0.25, grade: true, solution: false -->
```python
assert q10_1 == q10_4 or q10_2 == q10_3
```

<!-- grade_id: cell-27154033592c437f, points: 0.25, grade: true, solution: false -->
```python
assert q10_1 + q10_2 == q10_4 + q10_3
```

### Great work!